# <span style="color:teal;">CIS 211 Project 8:  &nbsp; Bears and Fish</span>

##### Due 11:00 P.M. Wednesday Mar 9

##### <span style="color:red">Group Members:</span>

Group work is allowed for this project.  Write the name and DuckID for each member of the group in this markdown cell:

Name:

Name:

Name:

Groups are exepcted to do more work.  See the sections at the end of this notebook for suggestions of how to extend the simulation.

##  <span style="color:teal;">Overview</span> 

The last project this term is a cellular automaton style biological simulation.  

The sytem is described in Chapter 11 of the Miller and Ranum text.  The simulation needs three class defintions:
* World defines a 2D grid where cells are accessed according to their row and column coordinates
* Fish defines one type of animal that can inhabit the world; fish can move and reproduce, and they can die from overcrowding
* Bear defines the other type of animal; bears can also move and reproduce, they eat fish, and they can die from starvation

**<span style="color:red">Hint:</span>** &nbsp; Before you start, think about what the Bear and Fish classes have in common, and consider defining a new class to use as the parent of both Bear and Fish.

##  <span style="color:teal;">Libraries</span> 

Your program will need to use `numpy` (a numeric processing library) and `random` (the builtin random number generator module).  You can import additional modules if you wish; add the import statements to this code cell.

```python
import numpy as np
import random
```

##  <span style="color:teal;">Event Log</span> 

A useful debugging technique is to save descriptions of events in a log file.  For this project we'll just print messages so they appear in the output area of the code cell that calls the top level simulator.

If you want to see messages printed as your program runs set `logging` to True in the code cell below, and include calls to `log` at the appropriate places in your methods.

**Note:** &nbsp; you can change the definition of the `log` function so it writes to a file or  prints additional information.

```python
logging = False

def log(message):
    if logging:
        print(message)
```

##  <span style="color:teal;">Classes, Functions, and Global Variables</span> 

As you write your code you can define additional classes, helper functions, and global variables.  Describe all of new classes, functions and variables in the following markdown cell, and put all of the definitions in one code cell right below the documentation.

[ your description here ]

```python
class Animal:
    '''
    Animal class is a parent which has methods that both the Fish and Bear classes share.
    '''
    def __init__(self, world, location, spec):
        self._world = world
        self._location = location
        self._spec = spec
        world.set_cell(location, spec)
        
    def location(self):
        return self._location
    
    def reproduce(self):
        pass
    
    def move(self):
        new_cell = random.choice(self._world.neighbors(self._location))
        if self._world.access_cell(new_cell) == None:
            self._world.move_cell(new_cell, self)
            self._world.delete(self._location)
            self._location = new_cell
        
```

##  <span style="color:teal;">World</span>

Describe your World class in the following markdown cell, and write the code for the class in the code cell below the markdown cell.

The minimum requiremens for this class are:
* the constructor should be passed the grid size (number of rows and columns)
* all cells should contain `None` when the world is initialized; later they can contain references to Fish or Bear objects
* override the methods that allow you to access a cell or store a value in a cell
* include a method named `biota` that returns a list of Fish and Bear objects currently in the grid

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Note:** &nbsp; There are no auto-grader tests for the World class; tests for this class will be done as part of the tests for the Bear and Fish classes.

<!-- grade_id: world_doc, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: world, points: 15, grade: true, solution: true -->
```python
class World:
    '''
    World class creates an array for the Fish and Bear classes to interact in
    '''
    def __init__(self, rows, columns):
        self._rows = rows
        self._columns = columns
        self._world = np.array([None] * (self._rows * self._columns)).reshape(rows, columns)
    
    def __repr__(self):
        return np.array_repr(self._world)
    
    def biota(self):
        animal_list = []
        for row in self._world:
            for col in row:
                if col != None:
                    animal_list.append(col)
        return animal_list
    
    def nrows(self):
        return self._rows
    
    def ncolumns(self):
        return self._columns
    
    def neighbors(self, location):
        r = (location[0], (location[1] + 1) % self.ncolumns())
        l = (location[0], (location[1] - 1) % self.ncolumns())
        u = ((location[0] - 1) % self.nrows(), location[1])
        d = ((location[0] + 1) % self.nrows(), location[1])
        dr = ((location[0] + 1) % self.nrows(), (location[1] + 1) % self.nrows())
        dl = ((location[0] + 1) % self.nrows(), (location[1] - 1) % self.nrows())
        ur = ((location[0] - 1) % self.nrows(), (location[1] + 1) % self.nrows())
        ul = ((location[0] - 1) % self.nrows(), (location[1] - 1) % self.nrows())
        return [r, l, u, d, dr, dl, ur, ul]
    
    def access_cell(self, location):
        return self._world[(location)]
    
    def set_cell(self, location, item):
        self._world.itemset((location), item)
        
    def move_cell(self, location, item):
        self._world.itemset((location), item)
        
    def cell_type(self, location):
        return type(self._world[(location)])
    
    def delete(self, location):
        self._world.itemset((location), None)
```

### <span style="color:teal;">Fish</span>

Use the markdown cell and code cell below to implement and describe your Fish class.

During the simulation a Fish object needs to behave as follows:

(1) Fish are susceptible to overcrowding:  if there are fish in 2 or more neighboring cells the fish dies (it's removed from the simulation)

(2) A fish can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new fish is placed in that cell

(3) A fish can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the fish moves there

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1 and 2 shown above
* `move` implements rule 3
* `location` returns the current grid location (row and column) of the object

Define a class variable named `breed_interval` to specify how many time steps a fish must be alive before it reproduces; the initial value for this variable is 12.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**NOTE:** &nbsp; The way the simulation is defined a fish might be eaten before the top level simulation calls the `live` method.  Your method should check to make sure the fish is still alive.  The easiest way to do this is to include an instance variable named `_alive` that is set to True when the fish is initialized and set to False when it dies.

<!-- grade_id: fish_doc, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: fish, points: 10, grade: true, solution: true -->
```python
class Fish(Animal):
    '''
    The Fish have a time step of 12 to breed.  Also the live method pertains to if they are too crowded.  
    '''
    
    breed_interval = 12
    
    def __init__(self, world, location):
        Animal.__init__(self, world, location, self)
        self._alive = True
        
    def __repr__(self):
        return "{0:^4}".format('FISH')
    
    def live(self):
        if self._alive == True:
            ncount = 0
            for cell in self._world.neighbors(self._location):
                neighbor = self._world.access_cell(cell)
                if type(neighbor) == Fish:
                    ncount += 1
            if ncount >= 2:
                self._world.delete(self._location)
                self._alive = False
                
            if Fish.breed_interval > 0:
                Fish.breed_interval -= 1
            if Fish.breed_interval == 0:
                new_cell = random.choice(self._world.neighbors(self._location))
                if self._world.access_cell(new_cell) == None:
                    Fish(self._world, new_cell)
                Fish.breed_interval = 12

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: fish_test_1, points: 3, grade: true, solution: false -->
```python
# A new world has no objects
w1 = World(5,5)
assert len(w1.biota()) == 0

# After adding a fish there should be one object
f1 = Fish(w1, (2,2))
assert len(w1.biota()) == 1

# Test the location method
assert f1.location() == (2,2)
```

<!-- grade_id: fish_test_2, points: 3, grade: true, solution: false -->
```python
# Setting breed_interval to 1 should cause a fish to reproduce when live is called
w2 = World(5,5)
Fish.breed_interval = 1
f2 = Fish(w2, (2,2))
f2.live()
assert len(w2.biota()) == 2

# Reset the interval to original value for remaining tests
Fish.breed_interval = 12
```

<!-- grade_id: fish_test_3, points: 3, grade: true, solution: false -->
```python
# Make three fish, the one in the middle should die from overcrowding
w3 = World(5,5)
f3 = Fish(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
f3.live()
assert len(w2.biota()) == 2
```

<!-- grade_id: fish_test_4, points: 3, grade: true, solution: false -->
```python
# When a fish moves it should be within one cell of its original location
w4 = World(5,5)
f4 = Fish(w4, (2,2))
f4.move()
r, c = f4.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

###  <span style="color:teal;">Bear</span>

Use the markdown cell and code cell below to implement and describe your Bear class.

During the simulation a Bear object needs to behave as follows:

(1) The bear looks for fish in each adjacent cell; if it finds one or more fish it eats one at random 

(2) If a bear has not eaten for certain number of time steps it dies (it's removed from the simulation)

(3) A bear can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new bear is placed in that cell

(4) A bear can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the bear moves there

Define a class variable named `breed_interval` to specify how many time steps a bear must be alive before it reproduces; the initial value for this variable is 8.  Define another class variable named `survive_without_food` to be the number of time steps a bear can live before it dies from starvation; the initial value for this variable is 10.

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1, 2, and 3 shown above
* `move` implements rule 4
* `location` returns the current grid location (row and column) of the object

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

<!-- grade_id: bear_doc, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: bear, points: 10, grade: true, solution: true -->
```python
class Bear(Animal):
    '''
    The Bear class has a time step of 8 to breed.  There is also an amount of steps they can live without eating a fish.  
    '''
    breed_interval = 8
    survive_without_food = 10
    
    def __init__(self, world, location):
        Animal.__init__(self, world, location, self)
        self._alive = True
        
    def __repr__(self):
        return "{0:^4}".format('BEAR')
    
    def live(self):
        if self._alive == True:
            for cell in self._world.neighbors(self._location):
                fish_list = []
                neighbor = self._world.access_cell(cell)
                if type(neighbor) == Fish:
                    fish_list.append(cell)
            if len(fish_list) > 0:
                target = random.choice(fish_list)
                self._world.delete(target)
                Bear.survive_without_food = 10
            else:
                Bear.survive_without_food -= 1
            if Bear.survive_without_food == 0:
                self._world.delete(self._location)
                self._alive = False
            if Bear.breed_interval > 0:
                Bear.breed_interval -= 1
            if Bear.breed_interval == 0:
                new_cell = random.choice(self._world.neighbors(self._location))
                if self._world.access_cell(new_cell) == None:
                    Bear(self._world, new_cell)
                Bear.breed_interval = 8
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: bear_test_1, points: 3, grade: true, solution: false -->
```python
# Test the Bear constructor and location method
w1 = World(5,5)
b1 = Bear(w1, (1,1))
assert len(w1.biota()) == 1
assert b1.location() == (1,1)
```

<!-- grade_id: bear_test_2, points: 3, grade: true, solution: false -->
```python
# Repeat the reproduction test for Bears
w2 = World(5,5)
Bear.breed_interval = 1
b2 = Bear(w2, (2,2))
b2.live()
assert len(w2.biota()) == 2
Bear.breed_interval = 8
```

<!-- grade_id: bear_test_3, points: 3, grade: true, solution: false -->
```python
# Make fish for the bear to eat, count the number of objects after eating
w3 = World(5,5)
b3 = Bear(w3, (2,2))
Fish(w3, (1,1))
Fish(w3, (3,3))
b3.live()
assert len(w2.biota()) == 2
```

<!-- grade_id: bear_test_4, points: 3, grade: true, solution: false -->
```python
# Setting the survival limit to 1 should cause a Bear to starve 
w4 = World(5,5)
Bear.survive_without_food = 1
b4 = Bear(w4, (2,2))
b4.live()
assert len(w4.biota()) == 0
```

<!-- grade_id: bear_test_5, points: 3, grade: true, solution: false -->
```python
# Repeat the move test for bears
w5 = World(5,5)
b5 = Bear(w5, (2,2))
b5.move()
r, c = b5.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

##  <span style="color:teal;">Top Level Functions</span>

Fill in the body of the `wbf` function so it returns a new World object with the specified number of rows and columns and with the specified number of Bear and Fish objects at random locations.

When we grade your project we will call `wbf` to make a World object and then use the main loop (implemented by `step_system`) to run the simulation.

<!-- grade_id: wbf, points: 5, grade: true, solution: true -->
```python
def wbf(nrows, ncols, nbears, nfish):
    '''
    Return a new World object with the specified number of Bear and Fish objects.
    '''
    world = World(nrows, ncols)
    w_list = []
    for i in range(nrows):
        for j in range(ncols):
            x = (i, j)
            w_list.append(x)
    for bear in range(nbears):
        random_cell = random.choice(w_list)
        Bear(world, random_cell)
        w_list.remove(random_cell)
    for bear in range(nfish):
        random_cell = random.choice(w_list)
        Fish(world, random_cell)
        w_list.remove(random_cell)
    return world
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: wbf_test, points: 3, grade: true, solution: false -->
```python
w = wbf(10,10,3,12)

dct = { Bear: 0, Fish: 0 }
for x in w.biota():
    dct[x.__class__] += 1

assert dct[Bear] == 3
assert dct[Fish] == 12
```

##### <span style="color:red">Experiments</span>

Run some experiments with the top level simulation loop and describe the results in the markdown cell below.  Some ideas of things to try:
* The settings for the Bear class `breed_interval` and `survive_without_food` variables come from the textbook.  Will the world ever run out of bears with these settings?
* Change the settings so `Bear.breed_interval` is larger than `Bear.survive_without_food`.  How does that change the outcome?
* Set the `breed_interval` counter for the Fish class to a smaller number (e.g. 4) so the world has more fish.  What effect does that have?
* Write a function that runs the simulation for a specified number of generations, or until there are no more objects left in the grid.  What combination of parameters leads to the largest number of time steps before the simulation halts.

```python
def step_system(world):
    for x in world.biota():
        x.live()
    for x in world.biota():
        x.move()
```

```python
x = wbf(7, 7, 1, 10)
for i in range(10):
    step_system(x)
    print(x)
    
```

Output:
```
array([[None, None, None, None, FISH, None, None],
       [None, None, FISH, None, FISH, None, None],
       [None, FISH, None, None, None, None, None],
       [None, None, None, FISH, None, None, None],
       [None, None, None, None, None, FISH, None],
       [BEAR, None, None, None, None, None, None],
       [None, None, None, None, FISH, None, None]], dtype=object)
array([[None, None, None, None, None, FISH, FISH],
       [None, None, FISH, FISH, None, None, None],
       [None, None, None, None, None, None, None],
       [None, None, FISH, None, None, None, FISH],
       [None, None, None, None, None, None, None],
       [None, None, None, None, None, None, None],
       [None, BEAR, None, None, FISH, None, None]], dtype=object)
array([[None, None, None, FISH, None, FISH, None],
       [None, None, FISH, None, None, None, None],
       [None, None, None, None, None, None, None],
       [FISH, None, None, None, None, None, FISH],
       [None, None, FISH, None, None, None, None],
 
```

<!-- grade_id: experiments, points: 10, grade: true, solution: true -->
Experiment results: Bears are pretty successful at eating the fish. Even with one bear, he seems to survive and reproduce failry easily.

##  <span style="color:teal;">Projects for Groups</span>

Here are two suggestions for ways to extend the simulation.  We will consider other types of extensions -- send a request to `conery@uoregon.edu` with your proposal.

Groups with three people should implement two extensions.  Groups with two people can choose either extension, or do both for extra credit.

* Implement the Plant class described in Section 10.7 of the textbook, and modify the Fish class so fish eat plants and die if they don't find enough food.

* Implement a GUI using `tkinter` that is similar to the Solar System GUI.  A canvas should display the world along with images for the Bear and Fish objects (you can download `Bear.gif` and `Fish.gif` from Canvas).  Use a spinbox or text entry box to specify the number of Bear and Fish objects and the number of time steps to run.  Include a Run button to start the simulation.


```python

```