# <span style="color:teal;">CIS 211 Project 1:  Hello, World</span>

##### Due 11:00 P.M. January 8, 2016

##### Reading:  none

This goal for this project is to work through the process for completing projects in CIS 211 this term:
* download a project description from Canvas, save it in a local `.ipynb` file
* open the file with Jupyter
* read the project description(s)
* add code and documentation to the notebook
* test your code
* e-mail the completed notebook to `cis211-project@cs.uoregon.edu`

###  <span style="color:teal">The `hello` Function</span>

Write a function named `hello` that will take a language name passed as a string argument and return a "hello, world" greeting in that language.

Examples:
<pre>
>>> hello('english')
'Hello, World!'

>>> hello('french')
'Bonjour, tout le monde!'

>>> hello('arabic')
'مرحبا يا عالم'
</pre>

Your function should be able to return a greeting in at least two languages -- English and French -- but you can have as many as you would like.

The strings to return for English and French are shown in the tests below.  Make sure you return these strings exactly as shown so your program passes the tests.

If the language passed as an argument is not one of the languages you implemented return None.
<pre>
>>> hello('valyrian')
</pre>

##### Style Points:

Try to write your program without using a bunch of `if-elif` cases. This is <span style="color:red"> not</span> a good way to structure the code in the body of your `hello` function:
<pre>
if lang == 'english':                   # bad style
    reply = 'Hello, World!'
elif lang == 'french':                  # ick
    reply = 'Bonjour, tout le monde!'
elif:                                   # ugh
    ...
</pre>

How do you want to handle situations where users pass a capitalized language name?  Do you want your function to return values for `hello("english")` and `hello("English")`?

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc, points: 30, grade: true, solution: true -->
This project translates the phrase "Hello, World!" to one of four desired languages (English, French, Zulu, and Chinese). The user input can include upper and lower case characters along with special characters. 

##### <span style="color:red">Code:</span>

<!-- grade_id: code, points: 0, grade: false, solution: true -->
```python
def hello(lang):
    '''This function takes an input language: lang and returns the english phrase "Hello, World!" translated to the desired language.'''
    word = ''
    for ch in lang:
        if ch.isalnum():
            word += ch
    word = word.lower()
    response = {'english' : 'Hello, World!', 'french' : 'Bonjour, tout le monde!', 'zulu' : 'Sawubona Mhlaba!', 'chinese' : '你好，世界！'}
    if word not in response:
        return None
    lang = response[word]
    
    return lang
```

##### <span style="color:red">Tests:</span>

Use the code cell below to test your `hello` function.  The human grader will evaluate these tests and assign style points based on the thoroughness of your tests.

<!-- grade_id: test, points: 20, grade: true, solution: true -->
```python
assert hello('ENGLISh') == 'Hello, World!'           #uppercase
```

```python
assert hello('french') == 'Bonjour, tout le monde!'  #base case
```

```python
assert hello('$$zulu!!') == 'Sawubona Mhlaba!'       #special characters
```

```python
assert hello('12345') == None                        #not translated
```

```python
assert hello('chinese') == '你好，世界！'             #non latin characters
```

##### <span style="color:red">Auto-Grader</span>

Your notebook will be processed by a program called `nbgrader`, which will run the tests below and automatically assign a grade.  You will earn credit for every cell that does not raise an Assertion error.

<!-- grade_id: english, points: 20, grade: true, solution: false -->
```python
assert hello('english') == 'Hello, World!'
```

<!-- grade_id: french, points: 20, grade: true, solution: false -->
```python
assert hello('french') == 'Bonjour, tout le monde!'
```

<!-- grade_id: klingon, points: 10, grade: true, solution: false -->
```python
assert hello('klingon') is None
```