# <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday Feb 5

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk
```

**Important:** &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

### <span style="color:teal">1. &nbsp; Hello, World (20 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe how the "hello" program works.

<!-- grade_id: doc_1, points: 5, grade: true, solution: true -->
The program makes a tk Label with the text "Hello, world!" and packs it into an application.

##### <span style="color:red">Code:</span>

Put your hello program in the following code cell.

<!-- grade_id: hello, points: 5, grade: true, solution: true -->
```python
hello_app = tk.Tk()
hello = tk.Label(hello_app, font = ('Helvetica'), text = 'Hello, world!')
hello.pack()
```

##### <span style="color:red">Autograder Tests:</span>

The three tests for this program check to see if the application has a single widget, that the widget is a label, and the label text starts with the string "hello".

**Do not edit, copy, or delete these cells.**

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 3, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 5, grade: true, solution: false -->
```python
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe your "payment" program.

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
This proagram creates an application with four labels, their cooresponding entry boxes, and a computing button. Invoking the button uses the first three box values to compute and show a payment on a loan depending on the amount, rate, and years of said loan.

##### <span style="color:red">Code:</span>

Put your payment program in the following code cell.

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
def monthly_payment(amt, pct, yrs):
    r = (float(pct) / 100) / 12
    p = 12 * float(yrs)
    pymt = (r * float(amt)) / (1 - ((1 + r) ** (-1 * p)))
    return pymt
    
def compute_payment_cb():

    amt = amount.get()
    pct = rate.get()
    yrs = years.get()
    
    payment.delete(0, tk.END)
    payment.insert(0, "${:.2f}".format(monthly_payment(amt, pct, yrs)))
    

payment_app = tk.Tk()

la_label = tk.Label(payment_app, text="Loan Amount:")
la_label.grid(row=0, column=0, sticky=tk.W, padx = 20, pady = 30)

ir_label = tk.Label(payment_app, text="Interest Rate:")
ir_label.grid(row=1, column=0, sticky=tk.W, padx = 20, pady = 30)

lp_label = tk.Label(payment_app, text="Loan Period:")
lp_label.grid(row=2, column=0, sticky=tk.W, padx = 20, pady = 30)

p_label = tk.Label(payment_app, text="Payment:")
p_label.grid(row=3, column=0, sticky=tk.W, padx = 20, pady = 30)

amount = tk.Entry(payment_app)
amount.grid(row=0, column=1, padx = 20)

rate = tk.Entry(payment_app)
rate.grid(row=1, column=1, padx = 20)

years = tk.Entry(payment_app)
years.grid(row=2, column=1, padx = 20)

payment = tk.Entry(payment_app)
payment.grid(row=3, column=1, padx = 20)

button = tk.Button(payment_app, text="Compute Payment", command=compute_payment_cb)
button.grid(row=4, column=0, columnspan=2, pady = 30)


```

##### <span style="color:red">Autograder Tests:</span>

The three tests for this program make sure the `monthly_payment` function computes the correct payment value, check to see if the GUI has the correct number of widgets, and test the GUI by inserting values into the entry boxes and clicking the button.

**Do not edit, copy, or delete these cells.**

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

### <span style="color:teal">3. &nbsp; Canvas (30 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Documentation:</span>

Use this markdown cell to describe how the "canvas" program works.

<!-- grade_id: doc_3, points: 5, grade: true, solution: true -->
This program creates a tk canvas and uses a the make_circle function to translate circle dimensions to draw circles on the canvas.

##### <span style="color:red">Code:</span>

Put your canvas program in the following code cell.

<!-- grade_id: canvas, points: 10, grade: true, solution: true -->
```python
def make_circle(canvas, radius, x, y, color):
    tlx = x - radius
    tly = y - radius
    brx = x + radius
    bry = y + radius
    canvas.create_oval(tlx, tly, brx, bry, fill=color)

canvas_app = tk.Tk()
canvas = tk.Canvas(canvas_app, width=400, height=400)
canvas.grid()
make_circle(canvas, 40, 100, 100, 'blue')
make_circle(canvas, 20, 200, 200, 'green')
make_circle(canvas, 80, 300, 300, 'yellow')


```

##### <span style="color:red">Autograder Tests:</span>

The tests for this program make sure the application has a single canvas widget and that the canvas has three circles.

**Do not edit, copy, or delete these cells.**

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```